# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets
#from sk learn library, import datasets modules

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```

```python
#classification --> predicting an outcome variable that is CATEGORICAL (vs continuous)
#regrission --> predicting an outcome variable that is CONTINUOUS (vs categorical)

#in sykit learn, you have to specify whether you wnt regression vs classification algorithm
#eg for tox outcomes, will want CLASSIFICATION (categorical)
#new H2O thing will choose for you
```

```python
#keyboard shortcuts to go from markdown to code: M = markdown; y = code
```

```python
#print(iris.DESCR)
#DESCR tells it put out text version of code

#iris.data
#gives you numerical output, including all features

#iris.feature_names
#gives you features list, which we will use to separate the sepcies of flow

#iris.target
#gives you numpy array *lots of 0s and 1s)
```

```python
#print(iris.DESCR)
#DESCR tells it put out text version of code

#iris.data
#gives you numerical output, including all features

iris.feature_names
#gives you features list, which we will use to separate the sepcies of flow

iris

iris.feature_names
iris.target
#assigned datasets module to the variable iris, and then ran "iris", which returns the dataset
```

```python
%who 
#tells you what things you have defined
#MAGIC COMMANDS --> start with a percent

%who_ls
#tells you.... list of variables

%ls

var_list = %who_ls

var_list
```

```python
#if something BREAKS in jupiter --> thing does not iterate execution count
#usually what happened is KERNAL got issue
#when you restart kernal, you LOSE ALL VARIABLES
#restart kernals with "00" then "enter"
```

```python

```

```python
from sklearn.svm import SVC
#svm is name of a module within the library sklearn
#SVC is a CLASS, which is a template for objects taht we're going to create (ie a blueprint)
X = iris.data
y = iris.target

classifier = SVC()
#putting parenthesis after SVC makes you CALL the class, which  gives you back a class instance --> instantiation of class
#assign the instance you get from SVC (a class), to the variable "classifier"
classifier.fit(X, y)
predictions = classifier.predict(X)
#this assigns predictions to a predictions variable
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-5-167bc204f964> in <module>
      2 #svm is name of a module within the library sklearn
      3 #SVC is a CLASS, which is a template for objects taht we're going to create (ie a blueprint)
----> 4 X = iris.data
      5 y = iris.target
      6 

NameError: name 'iris' is not defined
```

```python
X.shape
classifier.fit(X, y)
#this taught the computer something
```

Path to make shit happen --> standard workflow using sykit learn

1. Import
2. Instantiate (a class)
    -Estimators:
        -Classifiers
        -Regressors
3. Train test split (or CV)
4. Fit
5. Predict (and score)

-Overfitting (variance): tailored to a specific dataset (thus does not generalize well), misses the point, memorizes the data (eg complex polynomials)

-Underfitting (bias): accuracy could be better (eg maybe model is underpowered for problem); too general (eg simple linear model)

```python
predictions = classifier.predict(X)

predictions
#got 100% accuracy, b/c it learned data super well
#however, this is not very useful for predicting future data
#hence, need to TRAIN and then test fit on separate dataset
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-8-38d5dd97dcfc> in <module>
----> 1 predictions = classifier.predict(X)
      2 
      3 predictions
      4 #got 100% accuracy, b/c it learned data super well
      5 #however, this is not very useful for predicting future data

NameError: name 'classifier' is not defined
```

```python
predictions = classifier.predict(X)

predictions
#got 100% accuracy, b/c it learned data super well
#however, this is not very useful for predicting future data
#hence, need to TRAIN and then test fit on separate dataset
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-9-38d5dd97dcfc> in <module>
----> 1 predictions = classifier.predict(X)
      2 
      3 predictions
      4 #got 100% accuracy, b/c it learned data super well
      5 #however, this is not very useful for predicting future data

NameError: name 'classifier' is not defined
```

```python
from sklearn.svm import SVC
from sklearn.model_selection import train_test_split

X = iris.data
y = iris.target

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.33, random_state=42)

X_train
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-10-5b7219eeb24d> in <module>
      2 from sklearn.model_selection import train_test_split
      3 
----> 4 X = iris.data
      5 y = iris.target
      6 

NameError: name 'iris' is not defined
```

```python
from sklearn.svm import SVC
from sklearn.model_selection import train_test_split

X = iris.data
y= iris.target

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.33, random_state=42)
#random state is for reproducibility; used as "seed" # to send down path

#train_size : float, int, or None, (default=None)
    #If float, should be between 0.0 and 1.0 and represent the
    #proportion of the dataset to include in the train split. If
    #int, represents the absolute number of train samples. If None,
    #the value is automatically set to the complement of the test size.
    
    
X_train.shape
#matrix, 2 dimensions. Length = #rows, width = # columns
y_train.shape
#vector, just 1 column w/ 1 dimension hence only 1 #. It has 100 lengthclassifier = SVC()


X.shape

classifier.fit(X_train, y_train)
predictions = classifier.predict(X_test)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-11-e53aec6ce8ef> in <module>
      2 from sklearn.model_selection import train_test_split
      3 
----> 4 X = iris.data
      5 y= iris.target
      6 

NameError: name 'iris' is not defined
```