# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

iris= datasets.load_iris()
#from sk learn library, import datasets modules

# Load the iris dataset
# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)

```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
#from sklearn.svm import SVC

# YOUR CODE HERE
#raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```

```python
#classification --> predicting an outcome variable that is CATEGORICAL (vs continuous)
#regrission --> predicting an outcome variable that is CONTINUOUS (vs categorical)

#in sykit learn, you have to specify whether you wnt regression vs classification algorithm
#eg for tox outcomes, will want CLASSIFICATION (categorical)
#new H2O thing will choose for you
```

```python
#keyboard shortcuts to go from markdown to code: M = markdown; y = code
```

```python
#print(iris.DESCR)
#DESCR tells it put out text version of code

#iris.data
#gives you numerical output, including all features

#iris.feature_names
#gives you features list, which we will use to separate the sepcies of flow

#iris.target
#gives you numpy array *lots of 0s and 1s)
```

```python
#print(iris.DESCR)
#DESCR tells it put out text version of code

#iris.data
#gives you numerical output, including all features

iris.feature_names
#gives you features list, which we will use to separate the sepcies of flow

iris

iris.feature_names
iris.target
#assigned datasets module to the variable iris, and then ran "iris", which returns the dataset
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
%who 
#tells you what things you have defined
#MAGIC COMMANDS --> start with a percent

%who_ls
#tells you.... list of variables

%ls

var_list = %who_ls

var_list
```

Output:
```
SVC	 X	 classifier	 datasets	 iris	 predictions	 y	 
index.rst  Week1-assignment.ipynb
['SVC', 'X', 'classifier', 'datasets', 'iris', 'predictions', 'y']
```

```python
#if something BREAKS in jupiter --> thing does not iterate execution count
#usually what happened is KERNAL got issue
#when you restart kernal, you LOSE ALL VARIABLES
#restart kernals with "00" then "enter"
```

```python

```

```python
#from sklearn.svm import SVC
#svm is name of a module within the library sklearn
#SVC is a CLASS, which is a template for objects taht we're going to create (ie a blueprint)

#X = iris.data
#y = iris.target

#classifier = SVC()
#putting parenthesis after SVC makes you CALL the class, which  gives you back a class instance --> instantiation of class
#assign the instance you get from SVC (a class), to the variable "classifier"
#classifier.fit(X, y)
#predictions = classifier.predict(X)
#this assigns predictions to a predictions variable
```

```python
X.shape
classifier.fit(X, y)
#this taught the computer something
```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
SVC(C=1.0, cache_size=200, class_weight=None, coef0=0.0,
    decision_function_shape='ovr', degree=3, gamma='auto_deprecated',
    kernel='rbf', max_iter=-1, probability=False, random_state=None,
    shrinking=True, tol=0.001, verbose=False)
```

Path to make shit happen --> standard workflow using sykit learn

1. Import
2. Instantiate (a class)
    -Estimators:
        -Classifiers
        -Regressors
3. Train test split (or CV)
4. Fit
5. Predict 
6. Evaluate

-Overfitting (variance): tailored to a specific dataset (thus does not generalize well), misses the point, memorizes the data (eg complex polynomials)

-Underfitting (bias): accuracy could be better (eg maybe model is underpowered for problem); too general (eg simple linear model)

```python
predictions = classifier.predict(X)

predictions
#got 100% accuracy, b/c it learned data super well
#however, this is not very useful for predicting future data
#hence, need to TRAIN and then test fit on separate dataset
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1, 1, 2, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

```python
from sklearn.svm import SVC
from sklearn.model_selection import train_test_split

X = iris.data
y = iris.target

classifier = SVC()

X.shape

classifier.fit(X,y)

classifier.score(X, y)
#Returns the mean accuracy on the given test data and labels.
```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
0.9866666666666667
```

```python
predictions = classifier.predict(X)

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.33, random_state=42)

X_train
```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
array([[5.7, 2.9, 4.2, 1.3],
       [7.6, 3. , 6.6, 2.1],
       [5.6, 3. , 4.5, 1.5],
       [5.1, 3.5, 1.4, 0.2],
       [7.7, 2.8, 6.7, 2. ],
       [5.8, 2.7, 4.1, 1. ],
       [5.2, 3.4, 1.4, 0.2],
       [5. , 3.5, 1.3, 0.3],
       [5.1, 3.8, 1.9, 0.4],
       [5. , 2. , 3.5, 1. ],
       [6.3, 2.7, 4.9, 1.8],
       [4.8, 3.4, 1.9, 0.2],
       [5. , 3. , 1.6, 0.2],
       [5.1, 3.3, 1.7, 0.5],
       [5.6, 2.7, 4.2, 1.3],
       [5.1, 3.4, 1.5, 0.2],
       [5.7, 3. , 4.2, 1.2],
       [7.7, 3.8, 6.7, 2.2],
       [4.6, 3.2, 1.4, 0.2],
       [6.2, 2.9, 4.3, 1.3],
       [5.7, 2.5, 5. , 2. ],
       [5.5, 4.2, 1.4, 0.2],
       [6. , 3. , 4.8, 1.8],
       [5.8, 2.7
```

```python
#from sklearn.svm import SVC
#from sklearn.model_selection import train_test_split

#X = iris.data
#y= iris.target

#X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.33, random_state=42)
#random state is for reproducibility; used as "seed" # to send down path

#train_size : float, int, or None, (default=None)
    #If float, should be between 0.0 and 1.0 and represent the
    #proportion of the dataset to include in the train split. If
    #int, represents the absolute number of train samples. If None,
    #the value is automatically set to the complement of the test size.
    
    
#X_train.shape
#matrix, 2 dimensions. Length = #rows, width = # columns
#y_train.shape
#vector, just 1 column w/ 1 dimension hence only 1 #. It has 100 lengthclassifier = SVC()


#X.shape

#classifier.fit(X_train, y_train)
#predictions = classifier.predict(X_test)
```

```python
predictions

X.shape
```

Output:
```
(150, 4)
```

```python

```